# Classes and Object Oriented Programming
## Color

```python
from nose.tools import assert_equal, assert_true, assert_false, assert_raises, assert_almost_equal
```

**Problem 1 (10 points):** Implement `validate_alpha`. The function takes an argument, alpha, and returns alpha as a floating point number. If alpha is not between 0 and 1, a `ValueError` is raised.



<!-- grade_id: cell-60eb0cfb3c74494f, points: 0, grade: false, solution: true -->
```python
def validate_alpha(alpha):
    alpha = float(alpha)
    if alpha < 0 or alpha > 1:
        raise ValueError("alpha must be between 0 and 1")
    return alpha
    
```

## Tests for Problem 1

<!-- grade_id: cell-6363ecb34040617b, points: 3, grade: true, solution: false -->
```python
assert_equal(validate_alpha("0.15"), 0.15)
```

<!-- grade_id: cell-d0598995b8b0768a, points: 2, grade: true, solution: false -->
```python
assert_equal(validate_alpha(0.15), 0.15)
```

<!-- grade_id: cell-a266fa5d6e4bcb82, points: 2, grade: true, solution: false -->
```python
assert_raises(TypeError, validate_alpha, [3.4])
```

<!-- grade_id: cell-8ea717d8a41a2b72, points: 3, grade: true, solution: false -->
```python
assert_raises(ValueError, validate_alpha, 3.4)
```

**Problem 2 (10 points):** Implement `validate_alpha`. The function takes an argument, alpha, and returns alpha as a floating point number. If alpha is not between 0 and 1, a `ValueError` is raised.

<!-- grade_id: cell-521bbc4e5b44df30, points: 0, grade: false, solution: true -->
```python
def validate_color(color):
    color = int(color)
    if color <0 or color > 255:
        raise ValueError("color must be between 0 and 255")
    return color
```

## Tests for Problem 2

<!-- grade_id: cell-ce5a0832c25a103a, points: 2, grade: true, solution: false -->
```python
assert_raises(TypeError, validate_color, [])
```

<!-- grade_id: cell-2c0476bdd9854f86, points: 2, grade: true, solution: false -->
```python
assert_raises(ValueError, validate_color, -5)
```

<!-- grade_id: cell-38422f742a9a6c64, points: 2, grade: true, solution: false -->
```python
assert_raises(ValueError, validate_color, 500)
```

<!-- grade_id: cell-e83b4831e56ff17a, points: 2, grade: true, solution: false -->
```python
assert_equal(validate_color(55), 55)
```

<!-- grade_id: cell-bde44a1ffca4b381, points: 2, grade: true, solution: false -->
```python
assert_equal(validate_color(47.3), 47)
```

**Problem 3. (30 Points):** Define a named tuple ``rgbalpha`` that represents an RGB$\alpha$ color. Define an `rgba` class that inherits from the ``rgbalpha`` named tuple and adds an attribute name. 

Because we are inheriting from an immutable class, we need to do our validation of the color using a [`__new__`](https://docs.python.org/3/reference/datamodel.html#object.__new__) method, which is a static method of the class not of the instance of the class.

Define the following methods:

* `invert_rgb`
    * See docstring for method behavior to implement
* `grayscale`
    * See docstring for method behavior to implement

* ``__str__``
    * the returned string should include the integer values as zero padded integers (e.g. `028`) for each color and the floating point value with 2 decimal (e.g. `0.27`) points for alpha
* ``__repr__``
    * The returned string should include the class name and at least the information in the `__str__` string. 
* ``__add__``
    * When adding two `rgba` instances (e.g. `c1` and `c2`) to create a new `rgba` instance (e.g. `c3`), the color channels of `c3`  should be the sum of the two values mod 256. For example, $c3=(c1.r+c2.r)\mod 256$.
    * The alpha value should be the maximum of the two alpha values ($c3.\text{alpha}=\max(c1.\text{alpha},c2.\text{alpha})$)
* ``__eq__``
    * In comparing equality, ignore the alpha values
* ``__abs__``
    * The absolute value should be the root mean square of the color values ignoring the alpha value.


<!-- grade_id: cell-f88642675bef5a77, points: 0, grade: false, solution: true -->
```python
from collections import namedtuple
import math
rgbalpha = namedtuple("rgbalpha",['r','g','b','alpha'])
class rgba(rgbalpha):
    def __new__(cls, *args, **kwargs):
        return super().__new__(cls, validate_color(args[0]),
                          validate_color(args[1]),
                          validate_color(args[2]),
                          validate_alpha(args[3]))
    def __init__(self, *args, name="null"):
        self.name = name

    @property
    def name(self):
        return self.__name
    
    @name.setter
    def name(self, value):
        self.__name = value.lower()
    
    def invert_rgb(self): 
        """
        this function inverts the RGB color by subtracting 
        all color values from 255 and returns a new rgba object with the 
        new color values.
        
        alpha is not modified
        
        name
        """
        return rgba(255-self.r, 255-self.g, 255-self.b, self.alpha, name=self.name+".inverted")
 
    def grayscale(self): 

        """
        this function converts RGB color to grayscale by using a 
        weight average formula: 0.299Red+0.587Green+0.114Blue
        """

        tmp = 0.299*self.r+0.587*self.g + 0.114*self.b
        return rgba(tmp, tmp, tmp, self.alpha, name=self.name+".gray")

    def __str__(self):
        return """(%03d, %03d, %03d, %3.2f)"""%(self.r,
                                                         self.g,
                                                        self.b,
                                                        self.alpha)
    def __repr__(self):
        return """rgba instance:  %s (red=%d, green=%d, blue=%d; alpha=%f)"""\
                %(self.name, self.r,self.g,self.b,self.alpha)

    def __add__(self,color):
        return rgba((self.r+color.r)%256,
                    (self.g+color.g)%256,
                    (self.b+color.b)%256,
                   max(self.alpha,color.alpha), name=self.name+"_"+color.name)

    def __eq__(self,color):
        return self.r == color.r and \
               self.g == color.g and \
               self.b == color.b
        
    def __abs__(self):
        return math.sqrt(self.r*self.r+self.g*self.g+self.b*self.b)
        
```

## Tests for Problem 3

<!-- grade_id: cell-e32110b72c2df921, points: 2, grade: true, solution: false -->
```python
c=rgba(100,32,47,0.9785)
assert_true('032' in c.__str__())
```

Output:
```
True
```

<!-- grade_id: cell-2369dbaef892ab7f, points: 3, grade: true, solution: false -->
```python
c=rgba(100,32,47,0.9785)
assert_true('0.9785' in c.__repr__())
```

<!-- grade_id: cell-451d146f57d41bec, points: 5, grade: true, solution: false -->
```python
c=rgba(100,32,47,0.9785)
cg = c.grayscale()
assert_equal(cg.r,54)
assert_equal(cg.g,54)
assert_equal(cg.b,54)
```

<!-- grade_id: cell-944110c1cf346841, points: 2, grade: true, solution: false -->
```python
c=rgba(100,32,47,0.9785)
assert_true(isinstance(c,tuple))
```

Output:
```
True
```

<!-- grade_id: cell-02b6b8b650c33e86, points: 3, grade: true, solution: false -->
```python
assert_true(isinstance(c,rgbalpha))
```

<!-- grade_id: cell-0759a552d24d6e5c, points: 5, grade: true, solution: false -->
```python
c1 = rgba(100,32,47,0.9785)
c2 =rgba(255,0,0,0.1)
c3 = c1+c2
assert_equal(c3.r,99)
assert_equal(c3.g,32)
assert_equal(c3.alpha,0.9785)
```

<!-- grade_id: cell-8a78cdf49f641d25, points: 3, grade: true, solution: false -->
```python
c1 = rgba(100,32,47,0.9785)
c2 =rgba(255,0,0,0.1)
assert_false(c1==c2)
```

<!-- grade_id: cell-f04832eef7f5f612, points: 2, grade: true, solution: false -->
```python
assert_true(c1==c1)

```

Output:
```
115.03477735015616
```

<!-- grade_id: cell-55d9fff0ce3b772b, points: 5, grade: true, solution: false -->
```python
assert_almost_equal(abs(c1),115.03477735015616)
```

```python

```